```python
# Initialize Otter
import otter
grader = otter.Notebook("lab3.ipynb")
```

![](img/571_lab_banner.png)

# Lab 3: Text classification and hyperparameter optimization

<br><br>

## Imports

```python
from hashlib import sha1
import numpy as np
import pandas as pd
from sklearn.dummy import DummyClassifier
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.model_selection import (
    GridSearchCV,
    RandomizedSearchCV,
    cross_validate,
    cross_val_score,
    train_test_split,
)
from sklearn.pipeline import make_pipeline
from sklearn.svm import SVC
```

<br><br>

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions <a name="si"></a>
rubric={mechanics}

You will receive marks for correctly submitting this assignment by following the instructions below:
    
- Be sure to adhere to the [general lab instructions](https://ubc-mds.github.io/resources_pages/general_lab_instructions/).
- [Here](https://github.com/UBC-MDS/public/tree/master/rubric) you can find the description of each rubric used in MDS.
- Make at least three commits in your lab's GitHub repository.    
- Push the final .ipynb file containing your solutions to your GitHub repository for this lab.        
- Before submitting your lab, run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Clear All Outputs` and then `Run -> Run All Cells`. Notebooks without the output displayed may not be graded at all (because we need to see the output in order to grade your work).  
- Ensure that you are enrolled in Gradescope through [Canvas](https://canvas.ubc.ca/courses/147081).
- Upload the .ipynb file to Gradescope and wait till your assignment to run on Gradescope.
- Verify that your plots/output are properly rendered in Gradescope.    
- If the .ipynb file is too large or doesn't render on Gradescope for any reason, also upload a pdf (preferably WebPDF) or html export of .ipynb file with your solutions so that TAs can view your submission on Gradescope. 
- Do not push or upload the data you download for this lab to your repository or Gradescope. (There is also a `.gitignore` in the repo to prevent this.)
- Below this cell, include a clickable link to your GitHub repository for the lab just below this cell.
</div>    

_Points:_ 2

YOUR REPO LINK GOES HERE

<!-- END QUESTION -->

<br><br>

## Introduction <a name="in"></a>
<hr>

In this lab, we'll focus on two things:
1. Working with text data
2. Hyperparameter optimization

As this is a quiz week, I'm skipping the worksheet section. 

## Exercise 1: Introducing the dataset and EDA <a name="1"></a>
<hr>

Let's develop our own SMS spam classification system using Kaggle's [SMS Spam Collection Dataset](https://www.kaggle.com/uciml/sms-spam-collection-dataset). The goal is to predict whether a specific text message is spam or not. 

Please download the CSV and save it as `spam.csv` locally under the `data` folder. 

> **Some text messages in this dataset may contain offensive language, a reflection of the realities on such platforms 😔. If you are sensitive to such content, please avoid reading the raw messages.** 


The starter code provided below assumes that the data file is saved as "spam.csv" and is located in the data directory of this lab folder.

```python
sms_df = pd.read_csv("data/spam.csv", encoding="latin-1")
sms_df = sms_df.drop(columns=["Unnamed: 2", "Unnamed: 3", "Unnamed: 4"])
sms_df = sms_df.rename(columns={"v1": "target", "v2": "sms"})
sms_df.head()
```

<br><br>

### 1.1 Data splitting 
rubric={autograde}

**Your tasks:**

1. Split `sms_df` into train (80%) and test (20%) splits with `random_state=123`. 

<div class="alert alert-warning">

Solution_1.1
    
</div>

_Points:_ 1

```python
train_df, test_df = None, None

...
```

```python
grader.check("q1.1")
```

<br><br>

### 1.2 Simple EDA 
rubric={autograde}

Note that traditional EDA methods for tabular data don't directly apply to text data. In this exercise, we'll conduct basic EDA to understand the data better.

**Your tasks:**

1. Calculate the average length in character count (e.g., the message 'OK' has a length of 2) for both spam and ham (non-spam) messages. Then, store the results in the appropriate variables given below.
2. Record the lengths (i.e., character counts) of the shortest and longest messages for both spam and ham in the corresponding variables provided below. 

<div class="alert alert-warning">

Solution_1.2
    
</div>

_Points:_ 3

```python
avg_spam_len = None
avg_ham_len = None

...
```

```python
print("Average spam length: {:.3f}".format(avg_spam_len))
print("Average spam length: {:.3f}".format(avg_ham_len))
```

```python
shortest_spam_len = None
longest_spam_len = None
shortest_ham_len = None
longest_ham_len = None

...
```

```python
print(f"Shortest spam length: {shortest_spam_len}")
print(f"Longest spam length: {longest_spam_len}")
print(f"Shortest ham length: {shortest_ham_len}")
print(f"Longest ham length: {longest_ham_len}")
```

```python
grader.check("q1.2")
```

<br><br>

<!-- BEGIN QUESTION -->

### 1.3 Discussion 
rubric={reasoning}

**Your tasks:**

1. Discuss any observed differences in the average lengths of spam and ham messages.
2. Show the shortest and longest spam and ham messages. If multiple messages share the same length, you can choose any to showcase.
3. If we want to train supervised machine learning models on this data, how would you recommend encoding the `sms` column? Would one-hot encoding be appropriate for this column? Explain your reasoning.

<div class="alert alert-warning">

Solution_1.3
    
</div>

_Points:_ 3

_Type your answer here, replacing this text._

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 1.4 Word clouds
rubric={viz}

**Your tasks:**

1. Generate two word clouds: one for spam messages and another for ham (non-spam) messages using [the `wordcloud` package](https://github.com/amueller/word_cloud). To install it in your course's conda environment, use: ```conda install -c conda-forge wordcloud```
3. Briefly discuss your observations. 

> Avoid installing packages directly from the notebook. If you leave the installation command within the notebook, Gradescope might encounter issues processing your autograded questions.

> This package wasn't covered in our lectures. Consider checking the relevant documentation to learn how to create the word clouds.

<div class="alert alert-warning">

Solution_1.4
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 2: Model building
<hr>

## 2.1 `DummyClassifier`
rubric={autograde}

Now it's time to build models. Let's start with a baseline model. Before proceeding, remember to create `X_train`, `y_train`, `X_test`, `y_test` from `train_df` and `test_df` above.

**Your tasks:**
1. Calculate the mean cross-validation score of the `DummyClassifier` using default hyperparameters. Store the result in the `mean_cv_score` variable.

<div class="alert alert-warning">

Solution_2.1
    
</div>

_Points:_ 1

```python
X_train, y_train = None, None
X_test, y_test = None, None
mean_cv_score = None

...
```

```python
grader.check("q2.1")
```

<br><br>

### 2.2 Running a cross-validation fold without sklearn tools 
rubric={autograde}

Sklearn provides a lot of useful tools like `make_pipeline` and `cross_validate`, which are awesome. But with these fancy tools it's also easy to lose track of what is actually happening under the hood. 

**Your tasks:**

1. Calculate the validation score for SVC on the first fold, meaning you should train on 80% and validate on the remaining 20% of the `X_train` data. Do this without using sklearn's `Pipeline`, `cross_validate`, or `cross_val_score`. Save the score from this fold in a variable named `fold_score`. Remember that cross_validation in sklearn doesn't shuffle data by default.
To begin, use the provided `CountVectorizer` and `SVC` objects, along with `X_train` and `y_train`. You'll need to further divide these using `train_test_split` with the parameter `shuffle=False`.

```python
countvec_q_2_2 = CountVectorizer(stop_words="english")
svc_q_2_2 = SVC(random_state=123)
```

<div class="alert alert-warning">

Solution_2.2
    
</div>

_Points:_ 4

```python
fold_score = None

...
```

```python
grader.check("q2.2")
```

<br><br>

### 2.3 `SVC` with default hyperparameters
rubric={autograde}

Now that we know how each cross-validation fold works, let's go back to our usual sklearn tools. 

**Your tasks:**

1. Build a pipeline for feature extraction using `CountVectorizer` with `stop_words='english'` and `SVC` with `random_state=123` and default hyperparameters. Report the mean cross-validation score.

<div class="alert alert-warning">

Solution_2.3
    
</div>

_Points:_ 1

```python
svc_pipe = None
mean_cv_svc = None

...
```

```python
grader.check("q2.3")
```

<br><br>

## Exercise 3: Hyperparameter optimization
<hr>

Seems like SVC is a promising model on this problem. So let's carry out hyperparameter optimization for SVC. For all the sub-parts of this question, we will be using our `svc_pipe` from the previous exercise.

So far we have been writing loops to try a bunch of different hyperparameter values and pick the one with lowest cross-validation error. This operation is so common that `scikit-learn` has some [built-in classes](https://scikit-learn.org/stable/modules/grid_search.html) to do it for you. In this exercise, we will explore two such classes:

1. [`sklearn.model_selection.GridSearchCV`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) 
2. [`sklearn.model_selection.RandomizedSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.RandomizedSearchCV.html)

<br><br>

<!-- BEGIN QUESTION -->

### 3.1 Optimizing `C` for SVM RBF
rubric={reasoning}

The code below carries out hyperparameter search over `C` by sweeping the hyperparameter through the values $10^{-2}, 10^{-1}, 1, 10^{1}, 10^{2}$ using [`GridSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) and 5-fold cross-validation and plots training and cross-validation curves (mean train and mean cross-validation scores for different values of `C` from the `param_grid`) with `C` on the x-axis (use log scale) and scores on the y-axis. 

**Your tasks:**

1. Identify the "optimal" hyperparameter value and its associated best cross-validation score from the search. Save these values in the specified variables.
2. Contrast this score with the one you obtained in exercise 2.3, where you used the SVC with default settings and reflect on whether selecting the parameter value with highest cross-validation score in this scenario is a reasonable choice.
   
A few tips about `GridSearchCV`: 

- The starter code below defines the parameter grid for `C` which you can pass to your `GridSearchCV`. Note the syntax `svc__C`. We have two steps in our pipeline and we can access the parameters of these steps using `__` to go deeper. So `svc__C` means `C` of `svc` step of the pipeline. 
- Setting `n_jobs=-1` should speed things up (if you have a multi-core processor).
- Similar to `cross_validate`, you can pass `return_train_score=True` to your `GridSearchCV` object.
- By default, after carrying out the search, `GridSearchCV` also refits a model with the entire training set with the best hyperparameter values. 

```python
param_grid = {"svc__C": 10.0 ** np.arange(-2, 3)}
```

```python
import matplotlib.pyplot as plt
def plot_train_cv(param_vals, train_scores, cv_scores, xlabel, ylabel="accuracy"):
    plt.semilogx(param_vals, train_scores, label="train")
    plt.semilogx(param_vals, cv_scores, label="valid")
    plt.legend()

    plt.xlabel(xlabel)
    plt.ylabel(ylabel)
```

```python
print(
    "We are carrying out %d 5-fold cross-validation experiments"
    % (np.prod(list(map(len, param_grid.values()))))
)

C_search = GridSearchCV(
    svc_pipe, param_grid, n_jobs=-1, cv=5, return_train_score=True
)
C_search.fit(X_train, y_train)
```

```python
train_scores = C_search.cv_results_["mean_train_score"]
cv_scores = C_search.cv_results_["mean_test_score"]
plot_train_cv(param_grid["svc__C"], train_scores, cv_scores, xlabel="C")
```

<div class="alert alert-warning">

Solution_3.1
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

```python
best_C = None
best_C_score = None

...
```

<!-- END QUESTION -->

<br><br>

### 3.2 Jointly optimizing `C` and `gamma`
rubric={autograde}

Now let's optimize `C` hyperparameter along with `gamma`. 

**Your tasks:**

1. Expand your grid search to cover the `gamma` hyperparameter in addition to `C`, sweeping the hyperparameter through values $10^{-2}, 10^{-1}, \ldots, 10^{2}$ for both `C` and `gamma`. Store the information in the appropriate variables below. 

<div class="alert alert-warning">

Solution_3.2
    
</div>

_Points:_ 4

```python
param_grid = None
gamma_C_search = None
best_gamma_C = None
best_gamma_C_score = None

...

print("Best hyperparameter values: ", best_gamma_C)
print("Best score: %0.3f" % best_gamma_C_score)
```

```python
grader.check("q3.2")
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.3 Discussion
rubric={reasoning}

**Your tasks:**

1. Did the best `C` value you found in this search match what you got when only optimizing `C` in section 3.1? In a broader context, should these values necessarily align? Explain your reasoning.

<div class="alert alert-warning">

Solution_3.3
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.4 Jointly optimizing SVM RBF and `CountVectorizer` hyperparameters with `RandomizedSearchCV`
rubric={accuracy,quality}

Like estimators, transformers also have hyperparameters. In this exercise, you'll jointly optimize hyperparameters of `SVC` as well as `CountVectorizer` using 
`RandomizedSearchCV`, which, as its name implies, tries random hyperparameter configurations instead of performing an exhaustive grid search. 

**Your tasks:**

1. Jointly optimize `C` and `gamma` hyperparameters for the SVC with an RBF kernel, as well as the `binary` and `max_features` hyperparameters for the `CountVectorizer`. Use [`RandomizedSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.RandomizedSearchCV.html) for this and display the best hyperparameter values and best scores found from your search. When selecting the hyperparameters, opt for reasonable value ranges or distributions, based on your discretion. You can use either a parameter grid or suitable distributions. 

<div class="alert alert-warning">

Solution_3.4
    
</div>

_Points:_ 5

```python
param_dist = None

...
```

```python
random_search = None

...
```

```python
best_params = None

...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.5  `RandomizedSearchCV` and `GridSearchCV`
rubric={reasoning}

**Your tasks:**

1. Name a situation in which random search would be strongly preferable over a grid search, and provide a brief explanation for this preference. 

<div class="alert alert-warning">

Solution_3.5
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 4: Test results <a name="5"></a>
<hr>

Now that we have done extensive hyperparameter search, it's time to try our best model on the test split. 

### 4.1 Report test scores
rubric={autograde}

**Your tasks:**
1. Evaluate your top-performing model on both the training and test sets, then assign the results to the respective variables provided below.

<div class="alert alert-warning">

Solution_4.1
    
</div>

_Points:_ 1

```python
best_cv_score = None
train_score = None
test_score = None

...

print("CV score: %0.3f" % best_cv_score)
print("Train score: %0.3f" % train_score)
print("Test score: %0.3f" % test_score)
```

<br><br>

<!-- BEGIN QUESTION -->

### 4.2 Discussion and evaluation
rubric={reasoning}

**Your tasks:**

1. How does your test accuracy compare to your validation accuracy? Briefly comment. 
2. Try out your best model on a couple of fake spam and ham (non-spam) text messages. You can start with some sample messages given below.  
3. Do you believe that this model is likely to identify spam messages in the wild with the observed accuracy? Briefly discuss. 

<div class="alert alert-warning">

Solution_4.2
    
</div>

_Points:_ 3

```python
sample_text_msgs = ["""Call 8375467843 immediately. URGENT!! 
                    You have a free gift waiting for you...""", # potential spam
                    "MDS is fun.", # potential non-spam
                    """This is very URGENT!! This week the MDS team has decided that 
                      all of you get 100% on all assignments and quizzes for free. 
                      Visit https://bit.ly/2T15k6V to check your grades.""" # potential spam                    
]
```

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 5: Food for thought
<hr>

Each lab will have a few challenging questions. In some labs, I will be including challenging questions which lead to the material in the upcoming week. These are usually low-risk questions and will contribute to maximum 5% of the lab grade. The main purpose here is to challenge yourself or dig deeper in a particular area. When you start working on labs, attempt all other questions before moving to these questions. If you are running out of time, please skip these questions. 

We will be more strict with the marking of these questions. There might not be model answers. If you want to get full points in these questions, your answers need to
- be thorough, thoughtful, and well-written
- provide convincing justification and appropriate evidence for the claims you make 
- impress the reader of your lab with your understanding of the material, your analytical and critical reasoning skills, and your ability to think on your own

![](img/eva-game-on.png)

<!-- BEGIN QUESTION -->

### (Challenging) 5.1
rubric={reasoning}

In the EDA we noticed that the average length of spam vs. non-spam messages is quite different yet we are not using this feature in our current model. In this exercise, you will incorporate this feature in the model. 

**Your tasks:**
1. Add a new numeric feature called `sms_length` in the modeling which represents the length of text in characters, along with the bag of words features. 
2. What's the cross-validation accuracy of SVC with default hyperparameters and only `sms_length` feature? 
3. What's the cross-validation accuracy of SVC with default hyperparameters and `sms_length` and `CountVectorizer` features together? 
4. Discuss your observations. 

<div class="alert alert-warning">
    
Solution_5.1
    
</div>

_Points:_ 1

_Type your answer here, replacing this text._

```python
...
```

```python
...
```

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### (Challenging) 5.2
rubric={reasoning}

**Your tasks:**

There are all sorts of software packages that make hyperparameter tuning with `scikit-learn` even more automated. For example:

- [hyperopt-sklearn](https://github.com/hyperopt/hyperopt-sklearn)
- [auto-sklearn](https://github.com/automl/auto-sklearn)
- [SigOptSearchCV](https://sigopt.com/docs/overview/scikit_learn)

Give one of these a try on the task above and report your thoughts. Or, if you're even more adventurous, you could try a package that isn't tied to `scikit-learn`. There are many options for you to play around with in your ample free time:

- [TPOT](https://github.com/rhiever/tpot)
- [hyperopt](https://github.com/hyperopt/hyperopt)
- [hyperband](https://github.com/zygmuntz/hyperband)
- [SMAC](http://www.cs.ubc.ca/labs/beta/Projects/SMAC/)
- [MOE](https://github.com/Yelp/MOE)
- [pybo](https://github.com/mwhoffman/pybo)
- [spearmint](https://github.com/HIPS/Spearmint)
- [BayesOpt](https://github.com/rmcantin/bayesopt)

Note: this list isn't meant to be exhaustive. 

In other news, the [Amazon SageMaker](https://aws.amazon.com/sagemaker/) is also supposed to do hyperparameter optimization for you (among many other things it does).

<div class="alert alert-warning">
    
Solution_5.2
    
</div>

_Points:_ 1

<!-- END QUESTION -->

<br><br><br><br>

**Before submitting your assignment, please make sure you have followed all the instructions in the Submission Instructions section at the top.**

Congratulations on finishing the homework. You made it through a challenging week! Well done 👏👏! 

![](img/eva-well-done.png)



Tests for q1.1 (points: 1)
```
>>> assert not train_df is None and (not test_df is None), 'Please use the provided variables'
>>> assert train_df.shape == (4457, 2), 'The split is incorrect'
>>> assert test_df.shape == (1115, 2), 'The split is incorrect'
>>> assert list(train_df.loc[247]) == ['ham', 'Kallis wont bat in 2nd innings.'], 'Please check the random state'
>>> assert list(test_df.loc[3720]) == ['ham', 'Blank is Blank. But wat is blank? Lol'], 'Please check the random state'
```


Tests for q1.2 (points: [1, 1, 1])
```
>>> assert sha1(round(avg_spam_len, 3)).hexdigest() == 'c1a0dcae8c122b480a14e9f483cf6424817156c9', 'avg_spam_len seems incorrect'
>>> assert sha1(round(avg_ham_len, 3)).hexdigest() == 'a115649da1e6970e099667db61457df2389c616a', 'avg_ham_len seems incorrect'
```
```
>>> assert sha1(str(shortest_spam_len).encode('utf-8')).hexdigest() == 'bd307a3ec329e10a2cff8fb87480823da114f8f4', 'shortest_spam_len seems incorrect'
>>> assert sha1(str(longest_spam_len).encode('utf-8')).hexdigest() == 'bc15c774dca4499ea6fb42da7d216ca54f8c697e', 'longest_spam_len seems incorrect'
```
```
>>> assert sha1(str(shortest_ham_len).encode('utf-8')).hexdigest() == 'da4b9237bacccdf19c0760cab7aec4a8359010b0', 'shortest_ham_len seems incorrect'
>>> assert sha1(str(longest_ham_len).encode('utf-8')).hexdigest() == '4912f56aec6f0ac56bbb7fa9231e79891c48afc5', 'longest_ham_len seems incorrect'
```


Tests for q2.1 (points: 1)
```
>>> assert not X_train is None, 'Please use the provided variables'
>>> assert not y_train is None, 'Please use the provided variables'
>>> assert not X_test is None, 'Please use the provided variables'
>>> assert not y_test is None, 'Please use the provided variables'
>>> assert X_train.name == 'sms', 'Please check your X and y variables'
>>> assert X_test.name == 'sms', 'Please check your X and y variables'
>>> assert y_train.name == 'target', 'Please check your X and y variables'
>>> assert y_test.name == 'target', 'Please check your X and y variables'
>>> assert sha1(str(round(mean_cv_score, 3)).encode('utf-8')).hexdigest() == 'b5f478e5b18e752119bbb379d963c2357c5da422', 'mean_cross_val_score is incorrect.'
```


Tests for q2.2 (points: 4)
```
>>> assert not fold_score is None, 'Are you using the correct variable?'
>>> assert sha1(str(np.round(fold_score, 3)).encode('utf8')).hexdigest() == '5a18a87924a1faff708ddac91c052dbf60c7db78', "The fold_score doesn't look right."
```


Tests for q2.3 (points: 1)
```
>>> assert not svc_pipe is None, 'Please use the provided variable'
>>> assert sha1(str(round(mean_cv_svc, 3)).encode('utf-8')).hexdigest() == '4346e547194de927e67a1f4f4e2ba5381b4c5311', 'Are you storing mean cross-validation score in mean_cv_svc?'
```


Tests for q3.2 (points: [1, 2, 0.5, 0.5])
```
>>> assert not param_grid is None, 'Please use the provided variables'
>>> assert set(param_grid.keys()) == {'svc__C', 'svc__gamma'}, 'Please check the hyperparameters to optimize'
>>> assert 100.0 in param_grid['svc__gamma'], 'Please check the gamma values'
>>> assert np.prod(list(map(len, param_grid.values()))) == 25, 'The total number of experiments is incorrect'
```
```
>>> assert not gamma_C_search is None, 'Please use the provided variables'
>>> assert isinstance(gamma_C_search, GridSearchCV)
>>> assert len(list(gamma_C_search.get_params()['estimator'])) == 2, 'Please check your estimator and you may need to revisit 3.3'
```
```
>>> assert not best_gamma_C is None, 'Please use the provided variables'
>>> assert sha1(str(best_gamma_C.get('svc__C')).encode('utf8')).hexdigest() == 'c23f7da7070511444ebc75875fd9d202b5dd13cf', 'The best C value seems incorrect'
```
```
>>> assert not best_gamma_C_score is None, 'Please use the provided variables'
>>> assert np.isclose(round(best_gamma_C_score, 2), 0.98, atol=0.1), 'The score is incorrect'
```


Tests for q4.1 (points: 1)
